```python
import pandas as pd
import numpy as np
from pathlib import Path
import re
import requests
import time
```

```python
data = pd.read_csv('hemoglobin_ghana_clean.csv')
data
```

Output:
```
index  age sex  hemoglobin  city_country country
0       1   54   M        14.3           NaN   Ghana
1       2   42   M        15.1           NaN   Ghana
2       3   50   M        13.7           NaN   Ghana
3       4   34   M        14.4           NaN   Ghana
4       5   31   M        13.9           NaN   Ghana
..    ...  ...  ..         ...           ...     ...
65     66   22   M        11.9           NaN   Ghana
66     67   24   M        13.6           NaN   Ghana
67     68   57   M        12.9           NaN   Ghana
68     69   45   M        13.6           NaN   Ghana
69     70   55   F        12.0           NaN   Ghana

[70 rows x 6 columns]
```

```python
data = data.drop(columns = 'index')
data
```

Output:
```
age sex  hemoglobin  city_country country
0    54   M        14.3           NaN   Ghana
1    42   M        15.1           NaN   Ghana
2    50   M        13.7           NaN   Ghana
3    34   M        14.4           NaN   Ghana
4    31   M        13.9           NaN   Ghana
..  ...  ..         ...           ...     ...
65   22   M        11.9           NaN   Ghana
66   24   M        13.6           NaN   Ghana
67   57   M        12.9           NaN   Ghana
68   45   M        13.6           NaN   Ghana
69   55   F        12.0           NaN   Ghana

[70 rows x 5 columns]
```

```python
data.describe()
```

Output:
```
age  hemoglobin  city_country
count  70.000000   64.000000           0.0
mean   36.757143   13.959375           NaN
std    12.654192    1.209974           NaN
min    18.000000   11.600000           NaN
25%    28.000000   13.100000           NaN
50%    32.000000   14.000000           NaN
75%    46.750000   14.700000           NaN
max    86.000000   16.300000           NaN
```

```python
import plotly.express as px

```

```python
data['age'] = data['age'].replace('Null', np.nan)
data['age'] = pd.to_numeric(data['age'], errors='coerce')

# Drop rows with missing ages or hemoglobin values
data = data.dropna(subset=['age', 'hemoglobin'])

```

```python
scatter_fig = px.scatter(
    data,
    x='age',
    y='hemoglobin',
    color='sex',
    title='Hemoglobin Levels by Age and Sex',
    labels={'age': 'Age', 'hemoglobin': 'Hemoglobin (g/dL)', 'sex': 'Sex'},
)
scatter_fig.show()

```

```python
box_fig = px.box(
    data,
    x='sex',
    y='hemoglobin',
    title='Hemoglobin Distribution by Sex',
    labels={'sex': 'Sex', 'hemoglobin': 'Hemoglobin (g/dL)'},
    points='all'  # Shows all data points over the boxplot
)
box_fig.show()

```

```python
# Define age bins and labels
bins = [0, 12, 18, 40, 60, 100]
labels = ['Child', 'Teen', 'Young Adult', 'Middle Age', 'Senior']
data['age_group'] = pd.cut(data['age'], bins=bins, labels=labels)

# Group and calculate average hemoglobin
age_group_avg = data.groupby('age_group', observed=True)['hemoglobin'].mean().reset_index()

# Create bar plot
bar_fig = px.bar(
    age_group_avg,
    x='age_group',
    y='hemoglobin',
    title='Average Hemoglobin by Age Group',
    labels={'age_group': 'Age Group', 'hemoglobin': 'Avg Hemoglobin (g/dL)'}
)
bar_fig.show()

```

Output:
```
/var/folders/c2/z_0qd_k12mz2j28_b7t7zgth0000gn/T/ipykernel_12828/2949700821.py:4: SettingWithCopyWarning:


A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
```

```python

```

```python

```

Tests for q1 (points: None)
```
>>> isinstance(beowulf, str)
True
```
failure_message: check type to make sure the hidden tests will work
```
>>> # Why two possibilities?
>>> # The book changed since the last quarter we offered the project;
>>> # we're looking for the first output, but just in case the book changes again
>>> # while you're working on the project, we've provided the old expected output too.
>>> beowulf[:20] in ['\n\n\n\n\nBEOWULF\nAN ANGL', '\n\n\n\n\nProduced by Dav']
True
```
failure_message: make sure front matter not in string


Tests for q2 (points: None)
```
>>> 1000000 <= len(shakes) <= 1500000
True
```
failure_message: approx correct number of tokens
```
>>> elapsed <= 5
True
```
failure_message: shakespeare fast enough
```
>>> shakes[:3] == ['\x02', 'The', 'Complete'] and shakes[-3:] == ['William', 'Shakespeare', '\x03']
True
```
failure_message: check beginning and ending


Tests for q3 (points: None)
```
>>> bool((unif.mdl == 0.25).all())
True
```
failure_message: only one probability
```
>>> unif.mdl.shape[0] == 4
True
```
failure_message: number of tokens in mdl
```
>>> isinstance(unif.mdl, pd.Series)
True
```
failure_message: mdl correct type
```
>>> set(unif.mdl.index) == set('one two three four'.split())
True
```
failure_message: correct indices for mdl
```
>>> unif.probability(('five',)) == 0
True
```
failure_message: five not a token
```
>>> unif.probability(('one', 'two')) == 0.0625
True
```
failure_message: probability of given sequence appearing
```
>>> isinstance(unif.sample(1000), str)
True
```
failure_message: sample is string
```
>>> len(unif.sample(1000).split()) == 1000
True
```
failure_message: length of sample


Tests for q4 (points: None)
```
>>> unigram.mdl.nunique() == 3
True
```
failure_message: test mdl
```
>>> isinstance(unigram.mdl, pd.Series)
True
```
failure_message: mdl correct type
```
>>> set(unigram.mdl.index) == set('one two three four'.split())
True
```
failure_message: correct indices for mdl
```
>>> unigram.probability(('five',)) == 0
True
```
failure_message: five not a token
```
>>> isinstance(unigram.sample(1000), str)
True
```
failure_message: sample is string
```
>>> len(unigram.sample(1000).split()) == 1000
True
```
failure_message: sample length


Tests for q5 (points: None)
```
>>> isinstance(out[0], tuple)
True
```
failure_message: test create_ngrams
```
>>> out[0] == ('\x02', 'one') and out[2] == ('two', 'three')
True
```
failure_message: test create_ngrams
```
>>> set(bigrams.mdl.columns) == set('ngram n1gram prob'.split())
True
```
failure_message: test train
```
>>> bool(bigrams.mdl.shape == (6, 3) and bigrams.mdl['prob'].min() == 0.5)
True
```
failure_message: test train
```
>>> tokens = tuple('\x02 one two one three one two \x03'.split())
>>> bigrams2 = NGramLM(2, tokens)
>>> p = bigrams2.probability('two one three'.split())
>>> bool(np.isclose(p, (1/4) * (1/2) * (1/3), 0.01))
True
```
failure_message: test probability
```
>>> samp = bigrams.sample(3)
>>> len(samp.split()) == 4  # don't count the initial START token.
True
```
failure_message: test sample
```
>>> samp = bigrams.sample(3)
>>> samp[:2] == '\x02 ' and  set(samp.split()) <= {'\x02', '\x03', 'one', 'two', 'three', 'four'}
True
```
failure_message: test sample
```
>>> 36 <= len(out_5a1) <= 40 and all([len(x) == 2 for x in out_5a1]) and out_5a1[0] == ('\x02', 'Humpty')
True
```
failure_message: test sample
```
>>> #Older test cases (python 3.8)
>>> #np.isclose(out_5b1.groupby('n1gram').sum(), 1.0, 1.0).all() and (out_5b1['n1gram'].apply(len) == 1).all()
>>> bool(np.isclose(out_5b1.groupby('n1gram')['prob'].sum(),1.0).all() and (out_5b1['n1gram'].apply(len) == 1).all())
True
```
failure_message: test sample
```
>>> (out_5b1['ngram'].apply(len) == 2).all() and \
... bool(np.isclose(out_5b1.loc[out_5b1.ngram == ('Humpty', 'Dumpty'), 'prob'].squeeze(), 0.666666, atol=0.1))
True
```
failure_message: test sample
```
>>> bool(np.isclose(out_5c1.probability('Humpty Dumpty sat on a great fall'.split()), 0.01282051282051282, atol=0.1))
True
```
failure_message: test sample
```
>>> vc = pd.Series(out_5d1.split()).value_counts(normalize=True)
>>> {'\'', 'the', 'Humpty'}.intersection(set(vc.index[:6])) != {}
True
```
failure_message: most common sampled tokens
